# MLE for Poisson Distribution

In this notebook you will perform [maximum likelihood estimation](https://en.wikipedia.org/wiki/Maximum_likelihood), or MLE, for the parameter $\lambda$ in the [Poisson](https://en.wikipedia.org/wiki/Poisson_distribution) distribution and bootstrap the results.

## Imports

```python
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
from scipy.optimize import minimize
from scipy.stats import poisson
```

## Create a dataset

Generate a dataset from the Poisson distribution with a known parameter $\lambda$:

```python
np.random.seed(0)
λ = 3.0
data = np.random.poisson(λ, size=20)
```

## MLE with `scipy.optimize.minimize`

Use SciPy, write a function that performs MLE for the $\lambda$ parameter of the Poisson distribution. In your function you should:

* Use `scipy.stats.poisson` to calculate the negative log-likelihood of the Poisson distribution.
* Use `scipy.optimize.minimize` to minimize the negative log-likelihood.

<!-- grade_id: mle-poisson-f, points: 0, grade: false, solution: true -->
```python
def fit_lambda(data):
    """Perform MLE to estimate the λ parameter of the Poisson distribution.
    
    Parameters
    ----------
    data: ndarray
        The data to use in estimating lambda
    
    Returns
    -------
    lambda: float
        The MLE value for lambda
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: mle-poisson-g, points: 2, grade: true, solution: false -->
```python
assert abs(fit_lambda(3*np.ones(10))-3.0) < 1e-3
assert abs(fit_lambda(100*np.ones(100))-100.0) < 1e-3
assert abs(fit_lambda(data)-data.mean()) < 1e-3
```

Here is the estimated value of $\lambda$:

```python
fit_lambda(data)
```

## Bootstrap

We see that the MLE for $\hat{\lambda}$ is close to the original value of $\lambda=3.0$. Bootstrap resample this estimator 200 times (using your faster version) to find the distribution of $\hat{\lambda}$. Save the distributions of $\lambda$ values in an list or array named `estimates`:

<!-- grade_id: mle-poisson-j, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Plot the distibution of bootstrapped estimates using Matplotlib. Lable your axes, use a grid and customize the bins:

<!-- grade_id: mle-poisson-k, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Compute the print the 95% confidence interval:

<!-- grade_id: mle-poisson-l, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```