# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    same=True
    for i in range(0,10):
        if(list(x).count(str(i))!=list(y).count(str(i))):
            same=False
    return same
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
found=False
i=0
while(found==False):
    i=i+1
    found=(same_digits(str(i),str(2*i)) and same_digits(str(i),str(3*i)) and same_digits(str(i),str(4*i)) and same_digits(str(i),str(5*i)) and same_digits(str(i),str(6*i)))
print(i)
```

Output:
```
142857
```

```python

```