# Visualizing Subspots

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
```

## Read sunspot data

The data for this notebook comes from the "Yearly mean total sunspot number [1700 - now]" from the [SILSO](http://www.sidc.be/silso/datafiles) website. 

Use Pandas `read_csv` function to read the file `data/SN_y_tot_V2.0.csv` as a `DataFrame`:

* Look at the tests below for information on how to parse the CSV file.
* You will likely need to set the `sep`, `header`, `usecols` and `names` keyword arguments to `read_csv`.
* In the `SunspotsError` column, replace the missing value `-1` by `0`.
* Save your `DataFrame` in a variable `data`.

<!-- grade_id: mpl-sunspots-b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
data.head()
```

<!-- grade_id: mpl-sunspots-d, points: 2, grade: true, solution: false -->
```python
assert list(data.columns)==['Year', 'Sunspots', 'SunspotsError']
assert len(data)==317
assert data.Year.dtype==np.dtype(float)
assert data.Sunspots.dtype==np.dtype(float)
assert data.SunspotsError.dtype==np.dtype(float)
```

## Line chart

Make a line plot showing the sunspot count (y) as a function of year (x).

* Use the default Matplotlib 2.0 style as your starting point.
* Adjust the aspect ratio/size so that the steepest slope in your plot is *approximately* 1.
* Label the x- and y- axes and plot title.
* Use y axis gridlines.
* Pass data to `plt.plot` using the `data` keyword argument for Matpotlib.

<!-- grade_id: mpl-sunspots-e, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Measurements with error bars

Make a scatter plot with error bars showing the sunspot count (y) and its error (yerr) over time (x).

* Use the default Matplotlib 2.0 style as your starting point.
* Adjust the aspect ratio/size so that the steepest slope in your plot is *approximately* 1.
* Label the x- and y- axes and plot title.
* Use y axis gridlines.
* Pass data to `plt.errorbar` using the `data` keyword argument for Matpotlib.
* Filter the `DataFrame` to only include years after 1985.
* Use a format of `fmt='.k'`.

<!-- grade_id: mpl-sunspots-f, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Subplots by century

Now make 4 subplots, one for each century in the data set, with the x axis for time and the y axis for count. This approach works well for this dataset as it allows you to maintain mild slopes while limiting the overall width of the visualization. Perform similar customizations as above:

* Use the default Matplotlib 2.0 style as your starting point.
* Label the x- and y- axes and plot title.
* Use y axis gridlines.
* Pass data to `plt.plot` using the `data` keyword argument for Matpotlib.
* Set appropriate plot limits for each century.

<!-- grade_id: mpl-sunspots-g, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```