# Visualizing Exoplanets

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
```

## Exoplanet properties

Over the past few decades, astronomers have discovered thousands of extrasolar planets. The following paper describes the properties of some of these planets.

http://iopscience.iop.org/1402-4896/2008/T130/014001

Your job is to reproduce Figures 2 and 4 from this paper using an up-to-date dataset of extrasolar planets found on this GitHub repo:

https://github.com/OpenExoplanetCatalogue/open_exoplanet_catalogue

A text version of the dataset has already been put into this directory. The top of the file has documentation about each column of data:

```python
with open('data/open_exoplanet_catalogue.txt') as f:
    for i in range(50):
        print(f.readline(), end='')
```

Use `np.genfromtxt` with a delimiter of `','` to read the data into a NumPy array called `data`:

<!-- grade_id: mpl-exoplanet-a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mpl-exoplanet-b, points: 2, grade: true, solution: false -->
```python
assert data.shape==(1993,24)
```

Create a Pandas `Series` object named `masses` with the masses of the planets. Drop all missing values and filter out (exclude) all planet masses larger than 15.

<!-- grade_id: mpl-exoplanet-aa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mpl-exoplanet-bb, points: 2, grade: true, solution: false -->
```python
assert len(masses)==1041
assert type(masses)==pd.Series
assert len(masses[masses > 15])==0
```

Make a histogram of the distribution of planetary masses. This will reproduce Figure 2 in the original paper.

* Use the default Matplotlib 2.0 style as your starting point.
* Pick the number of bins for the histogram appropriately.
* Use y axis gridlines.
* Label the x- and y- axes and the plot title.
* Use appropriate x and y limits.

<!-- grade_id: mpl-exoplanet-c, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Make a scatter plot of the orbital eccentricity (y) versus the semimajor axis. This will reproduce Figure 4 of the original paper.

* Use the default Matplotlib 2.0 style as your starting point.
* Use a log scale on the x axis.
* Label the x- and y- axis and the plot title.
* Use appropriate x and y limits.
* Use black filled circles with an alpha of 0.3 for the marks.

<!-- grade_id: mpl-exoplanet-d, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```