# Matplotlib: Cars Dataset

In this notebook you will visualize the cars dataset using Matplotlib.

## Imports

```python
import numpy as np
import pandas as pd
from altair import load_dataset
```

```python
%matplotlib inline
from matplotlib import pyplot as plt
```

## Dataset

```python
cars = load_dataset('cars')
```

```python
cars.head()
```

## Scatterplot of displacement versus acceleration 

Use Matplotlib to create a scatterplot of displacement (y) versus acceleration versus (x):

* Use an alpha of 0.5
* Use the `data` keyword argument of Matplotlib to pass in the data.
* Label the x- and y- axes and plot title.
* Use x and y gridlines.

<!-- grade_id: mpl-cars-e, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Average horsepower by year

Create a new `DataFrame`, named `hp_year`, with two columns:

* The year of the date only (`YearOnly`).
* The average horsepower for each year.

<!-- grade_id: mpl-cars-f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
hp_year
```

<!-- grade_id: mpl-cars-h, points: 2, grade: true, solution: false -->
```python
assert list(hp_year.columns)==['YearOnly', 'Horsepower']
assert list(hp_year.YearOnly.unique())==[1970, 1971, 1972, 1973, 1974, 1975, 1976, 1977, 1978, 1979, 1980, 1982]
assert len(hp_year)==12
```

Use Matplotlib to create a line chart of average horsepower (y) versus year (x):

* Out both lines and points on your chart.
* Use the `data` keyword argument.
* Label the x- and y- axes and plot title.
* Use x and y axes gridlines.

<!-- grade_id: mpl-cars-i, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Displacement histogram

Use Matplotlib to create a histogram of the displacement:

* Use the `data` keyword argument.
* Set the number of bins appropriately.
Label the x- and y- axes and plot title.
* Use y axes gridlines.

<!-- grade_id: mpl-cars-j, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```