# Simpsons: Merging and Concatenation

## Imports

```python
import numpy as np
import pandas as pd
from pandas import DataFrame, Series
```

## Relationships between DataFrames

When you have multiple `DataFrame`s that have common keys you can have **relationships** between the entities in the different `DataFrame`s. There are three types of entity relationships that are possible:

* 1-to-1
* 1-to-many
* many-to-many

Here is a small data set from the TV show [The Simpsons](https://en.wikipedia.org/wiki/The_Simpsons) to illustrate these relationshps.

First, here is a `DataFrame` with students' first and last names, along with a unique student id:

```python
students = DataFrame({'fname': ['Bart','Lisa','Milhouse'],
                      'lname': ['Simpson','Simpson','Van Houten']},
                     index=list('abc'))
students
```

Here is a `DataFrame` with the student social security numbers, indexed by their unique student id:

```python
ssns = DataFrame({'ssn':[1234,5678,9101]}, index=list('abc'))
ssns
```

Each student can have aliases or nicknames:

```python
aliases = DataFrame({'alias':['Bartman','Bartron','Cosmos','Truth Teller','Lady Penelope Ariel',
                              'Jake Boyman','Lou La Trec','Eagle Eye','Maestro'],
                     'student': list('aaabbbccc')})
aliases
```

Here are the student home addresses:

```python
addresses = DataFrame({'address':['742 Evergreen Terrace','742 Evergreen Terrace','316 Pikeland Ave.']},
                      index=list('abc'))
addresses
```

A table of courses the students can be enrolled in:

```python
courses = DataFrame({'name':['Biology','Math','PE','Underwater electronics']}, index=range(4))
courses
```

This table contains the enrollment for each course. Each row of this table has a student and course.

```python
enroll = DataFrame({'student':['a','b','b','c','c','c']},index=(2,0,1,0,1,2))
enroll
```

## 1-1 relationships

* Each student has exactly one SSN.
* Each SSN belongs to exactly one student.

Create a `DataFrame` with the students' first name, last name and social security number:

<!-- grade_id: simpsons-g, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
merge1
```

<!-- grade_id: simpsons-i, points: 2, grade: true, solution: false -->
```python
assert list(merge1.columns)==['fname', 'lname', 'ssn']
assert list(merge1.index)==list('abc')
```

## 1-many relationships

### Students and addresses

* Each student has exactly one address.
* Each address can have many students.

Create a `DataFrame` with the students' first name, last name and address:

<!-- grade_id: simpsons-j, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
merge2
```

<!-- grade_id: simpsons-l, points: 2, grade: true, solution: false -->
```python
assert list(merge2.columns)==['fname', 'lname', 'address']
assert list(merge2.index)==list('abc')
```

### Students and aliases

* Each student can have many aliases.
* Each alias belong to exactly one student.

Create a `DataFrame` with the students' first name, last name and alias. The index of the data frame should be the student column (a, b, c).

<!-- grade_id: simpsons-m, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
merge3
```

<!-- grade_id: simpsons-o, points: 2, grade: true, solution: false -->
```python
assert list(merge3.columns)==['fname', 'lname', 'alias']
assert list(merge3.index)==list('aaabbbccc')
```

## Many-many relationships

* A student can take multiple classes.
* A single class can have multiple students.

Create a `DataFrame` with the students first name, last name and student (a, b, c) and course name. Multiple merges may be required.

<!-- grade_id: simpsons-p, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
merge4
```

<!-- grade_id: simpsons-r, points: 2, grade: true, solution: false -->
```python
assert list(merge4.columns)==['fname', 'lname', 'student', 'name']
assert len(merge4)==7
```

## Concatenation

Use Pandas' `concat` function to combining the `students` and `ssns` `DataFrame`s by columns

<!-- grade_id: simpsons-s, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
concat1
```

<!-- grade_id: simpsons-u, points: 2, grade: true, solution: false -->
```python
assert list(concat1.columns)==['fname', 'lname', 'ssn']
assert list(concat1.index)==list('abc')
```

Do the same thing for the `students`, `ssns` and `addresses` `DataFrame`s:

<!-- grade_id: simpsons-v, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
concat2
```

<!-- grade_id: simpsons-x, points: 2, grade: true, solution: false -->
```python
assert list(concat2.columns)==['fname', 'lname', 'ssn', 'address']
assert list(concat2.index)==list('abc')
```